---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## Insert exercise title here

```yaml
type: NormalExercise
key: c6ebe2b16b
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 93ba3a7759
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Insert exercise title here

```yaml
type: NormalExercise
key: 7847f2b825
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```
