---
title: 'Template Chapter 1'
description: 'This is a template chapter.'
---

## Exercise 1

```yaml
type: NormalExercise
key: h2oz7i0y6n
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---

## Exercise 2

```yaml
type: NormalExercise
key: 93ba3a7759
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Exercise 3

```yaml
type: NormalExercise
key: 7847f2b825
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```
